# Assessment 1

#### This course is assessed by two, individual, coursework exercises and the weekly quiz:
* Quiz (5%): issued each week to assess understanding of previous weeks topics
* Assignment 1 (25%): given out Friday 16th of Febuary & submitted via a dropbox in the Assessment area of the classes Learn page on Tuesday the 5th of March before 4pm
* Assignment 2 (70%): given out Friday 8th of March & submitted on Tuesday the 2nd of April

This notebook relates to _**Assessment 1**_

If your personal circumstances (e.g. health, work pressure, caring responsibilities, etc) result in you being unable
to complete the assignment by the due date it is essential you apply for an extension from:
https://www.ed.ac.uk/student-administration/extensions-special-circumstances

* It is expected that you discuss the assignment with tutors and lectures and ask for help with this exercise.  The more detailed the question the easier it is to help.
  * 	Please note that the University has a very strict plagiarism policy and several automated systems that check submitted work. Any submitted work must be your own but obviously, it is acceptable to base your code on examples from any sourc.

* More comments (in the code) and explanation in markdown cells) more marks!

* 
*It is acceptable to use Chatbots during the development and debugging of your solutions. However, 
after the submission date we may interview some, or all, the class to discuss the details of the code 
individuals have presented. You will be expected to be able to explain 
every line and answer questions about the impact small changes might have on the overall function..



#### This assessment asks you to do the following tasks for a dataset of two CSV files we supply.

0. Load the libraries (1 mark)
   
__Clean-up the data__
1. Import the two .csv files into the notebook and create dataframe (2 marks)
2. Perform initial inspection of the data (summary of the size and shape of the data) (8 marks)
3. Change text responses (y,yes,ok etc..) into 1s or 0s (4 marks)
4. Create a single dataframe to hold all the loaded data (2 marks)
   
__In the combined dataframe__

5. Replace any NaN or blank values with nominal values (given in row zero) (4 marks)
6. Convert all numbers to floats (some are integer or strings) (2 marks)
7. Delete any duplicate rows (4 marks)

__Analyse data by__

8. Change the values in the data set to show the error (i.e. value less the nominal value) (4 marks)
9. Add a new row to the dataframe that shows the average difference between nominal and actual values for each column (i.e. dimension) (4 marks)

__Visualise data__

10. Plot a histogram of the values calculated in step 8 (i.e. average variation of each dimension) (5 marks)
11. Investigate if there are correlations between the variation in dimensions (you might need to create a new dataframe to do this) (10 marks)
12. Visualise the results of the correlation investigation (10 marks)

__Discuss the Data__

13. Report some conclusions about the data (5 mark)

This assessment is based on the data cleaning, merging, analysis and visualisation of inspection data from a Coordinate Measuring Machine [CMM](https://en.wikipedia.org/wiki/Coordinate-measuring_machine). The data provided involves inspection data for a forged component provided in the form of two .csv files (**A1_CMM_DOE1**), and (_**A1_CMM_DOE2**_). The .csv files consists of values of 12 key dimensions and a visual inspection. The original data has been simplified to show the nominal value for each dimension, without a tolerance in the first row, However there are inconsistencies in data input, which need to be cleaned in order to analysised and visualize effectively. 

![Component](https://raw.githubusercontent.com/drewsherlock/dsim_workshops/main/Picture_dataset_structure.jpg)

<span style="color:red">

__Put your anwers in the cells containing the lines:__

_# YOUR CODE HERE_  
raise NotImplementedError()
 
__Delete the line:__  

raise NotImplementedError()

__Do not add or change the order of the cells__
    
</span>

### Step 0: 
Load the pandas and numpy libraries (as _**pd**_ and _**np**_ respectively)

<!-- grade_id: cell-fa6a0df4831a2ba3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-72492488daf5fcbc, points: 1, grade: true, solution: false -->
```python
# Run this cell to check if the libraries have been loaded 

import types

try:
    assert isinstance(pd, types.ModuleType), "pandas has not been imported."
    assert pd.__name__ == 'pandas', "pandas has not been imported as 'pd'."
    assert isinstance(np, types.ModuleType), "numpy has not been imported."
    assert np.__name__ == 'numpy', "numpy has not been imported as 'np'."
except (AssertionError, NameError) as e:
    print("Error in the loading of libraries:", str(e))
else:
    print("Well done libraries are loaded")
```

### Step 1: 
Import the two .csv files called A1_CMM_DOE1 and A1_CMM_DOE2 into two Pandas dataframes called _**data1**_ and _**data2**_
URLs of the files are: 
* 'https://raw.githubusercontent.com/drewsherlock/dsim_workshops/main/A1_CMM_DOE1.csv' and
* 
'https://raw.githubusercontent.com/drewsherlock/dsim_workshops/main/A1_CMM_DOE2.csv'

<!-- grade_id: cell-fb4eb9ffc9d882ae, points: 0, grade: false, solution: true -->
```python
# Import the two .csv files you have been given into two dataframes called data1 and data2 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-77e9334d360e645b, points: 2, grade: true, solution: false -->
```python
# Run this cell to test if the data has been correctly loaded 
# Note that the test is not checking which datasets have been loaded, only that something has been loaded

try:
    assert isinstance(data1, pd.DataFrame), "data1 not found. Make sure you've imported the CSV to a DataFrame named data1."
    assert isinstance(data2, pd.DataFrame), "data2 not found. Make sure you've imported the CSV to a DataFrame named data2."
except (AssertionError, NameError) as e:
    print("Error in the loading of CSV files:", str(e))
else:
    print("Variables data1 and data2 refer to dataframes")

```

### Step 2: 
Perform initial inspection of the data in data1 and data2 by: 

* Assign the number of rows (r) and columns (c) of each dataframe to the variables: _**data1_r, data1_c**_ & _**data2_r, data2_c**_ 

* Assign a dataframe containing the first 4 rows of the each dataset to the variable _**data1_h**_ and _**data2_h**_ .

* Create two lists with variable names _**data1_cols**_ and _**data2_cols**_  that contain the column labels of _**data1**_ and _**data2**_ 

* Summarizes the numerical values in the dataframes using the “describe” function and assign the re_**data1_des**_ and _**data2_des**_
    
                                                                                                        

<!-- grade_id: cell-14c64a2268bf073c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4fd2f67746959926, points: 8, grade: true, solution: false -->
```python
try:
# check the number of rows and columns are correct
    assert (data1_r, data1_c) == data1.shape, "data1_r and/or data1_c not found or incorrect"
    assert (data2_r, data2_c) == data2.shape, "data2_r and/or data2_c not found or incorrect"

# check heads of data1 and data2
    assert data1_h.equals(data1.head(4)), "Head of data1 is incorrect"
    assert data2_h.equals(data2.head(4)), "Head of data1 is incorrect"

# Check list of column names of each data frame to data1_cols and data2_cols
    assert data1_cols == data1.columns.tolist(), "data1 Coloumn list is incorrect"
    assert data2_cols == data2.columns.tolist(), "data2 Coloumn list is incorrect"

# Summarize the numerical values in the data frames using the "describe" function
    assert data1_des.equals(data1.describe()), "data1 description is wrong"
    assert data2_des.equals(data2.describe()), "data2 description is wrong"

except (AssertionError, NameError) as e:
    print("Error in inspecting the data:", str(e))
else:
    print("The content of the dataframes has been described")

```

### Step 3 : 

From each dataframe ( _**data1**_ and _**data2**_ ) change all the text strings in the last column to the 1 if they are "Y", "y", "Yes", "yes", "ok","okay" or "Okay" and 0 if they have the value "N", "n", "No", "no".\es())

<!-- grade_id: cell-7a32c4519ae73632, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bc9b16a6fa554f46, points: 4, grade: true, solution: false -->
```python
# Check Solution
try:
    # Check if all values in last column  are either 1 or 0
    data1condition1 = (data1.iloc[:, -1] == 1).all()
    data1condition2 = (data1.iloc[:, -1] == 0).all()
    assert (not data1condition1  | data1condition2), "Text values still exist in last column of data1."
    
    # Check if all values in last column  are either 1 or 0
    data2condition1 = (data2.iloc[:, -1] == 1).all()
    data2condition2 = (data2.iloc[:, -1] == 0).all()
    assert (not data2condition1 | data2condition2), "Text values still exist in last column of data1."

except (AssertionError, NameError) as e:
    print("Error the processing of the last column (text values remain:", str(e))
else:
    print("Well done text in both dataframes has been converted to 1s and 0s")
```

### Step 4 : 

Combine the two dataframes _**data1**_ and _**data2**_ and call the result _**data3**_ 

<!-- grade_id: cell-4c835031a49d019e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d857ce2ffe16d03, points: 2, grade: true, solution: false -->
```python
# Check Solution
try:
    assert data3.shape[0] == 101, "Number of rows does not match expected value (101)."
    assert data3.shape[1] == 14, "Number of columns does not match expected value (14)."
    
except (AssertionError, NameError) as e:
    print("Error data3 does not refer to a dataframe or shape is worng", str(e))
else:
    print("data3 has been created and has the right shape!")
```

### Step 5:  

Replace any _**NaN**_ and _**Blanks**_ in _**data3**_ with the __nominal value__ of the dimension

<!-- grade_id: cell-ea39b6bfe3765670, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-570993753bf0474b, points: 4, grade: true, solution: false -->
```python
# Check solution
try: 
    assert data3.iloc[20,2] == 42.0, "Test NaN value has been replaced."

except (AssertionError, NameError) as e:
    print("Still NaNs and blancks in dataframe data3", str(e))
else:
    print("Great, the blanks and NaNs have been filled")
```

### Step 6: 
Convert to numbers (i.e. floats) all the number values in _**data3**_  (recall some values look like numbers, but are actually string types).
See https://study.com/academy/lesson/data-types-in-programming-numbers-strings-and-others.html for data type introduction

<!-- grade_id: cell-b3f2b5b9f4a7225a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0a7ffd1624a3969e, points: 2, grade: true, solution: false -->
```python
# Check if all values are float data type
try:
    assert data3.dtypes.eq(float).all(), "data3 still contains values that are not floating point numbers."
    
except (AssertionError, NameError) as e:
    print("There are still some none numerical values in the data frame", str(e))
else:
    print("Hooray the data3 dataset is looking good with only numbers in it")

```

## Step 7: 
Delete any rows in _**data3**_ with exactly the same values in all the columns (i.e. duplicates)

<!-- grade_id: cell-1d7ac373c2517a03, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f19169fcbfee335, points: 4, grade: true, solution: false -->
```python
# Check Solution
try:
    assert not data3.duplicated().any(), "Duplicate rows or columns still exist in data3." 
    
except (AssertionError, NameError) as e:
    print("There are still duplicate rows in the data frame", str(e))
else:
    print("Hooray the data3 dataset is ready to be analysed")
```

### Step 8: 

_**Analysis 1**_:  Update the values in **data3** so they show the __difference__ between nominal and actual vaules 

<!-- grade_id: cell-adaa0e90fb695d27, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f76c6dc580098487, points: 4, grade: true, solution: false -->
```python
# Check Solution
try:
    assert data3.iloc[2,2] < 1.0,  "Values in data3 appear unchanged" 

except (AssertionError, NameError) as e:
    print("The numbers in data3 should be mostly less than 1", str(e))
else:
    print("Hooray data3 is showing the errors (not the actual values) ")

```

### Step 9: 

_**Analysis 2**_:  Add a new row to the dataframe that shows the _**average difference**_ between nominal and actual values, for each column (i.e. dimension)

<!-- grade_id: cell-ea49c3b1c70e38bf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6e0d432cf3240bd8, points: 4, grade: true, solution: false -->
```python
# Check Solution
try: 
    assert data3.shape[0] == 98, "The number of rows doesn look right"
    
except (AssertionError, NameError) as e:
    print("The number of rows in data3 has not changed", str(e))
else:
    print("Hooray data3 has a new row ")

```

### Step 10: 

Visualize the relative variability of the dimensions by plotting a histogram of the values in bottom row (the one which was added in the previous step). Don't include the values of the last column in the plot (e.g. ok, yes, NO, n, Okay etc)

<!-- grade_id: cell-d1b928ebc4a33075, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Step 11: 
**Analysis 3** Investigate if there are correlations between the size of variations (i.e. errors) in dimensions 
(you might need to create a new dataframe to do this)

<!-- grade_id: cell-1fad1ce69cca74fb, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Step 12:
Visualise the results of the correlation investigation

<!-- grade_id: cell-ae407b36377d4c09, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Step 13:

Suggest 3 conclusions that can be drawn from this dataset, write your conclusions in the cell below using the markdown language to format your text

<!-- grade_id: cell-1fc34e64e068fdab, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE